### Final Project

In this project, we're going to compare the performance of various models in predicting the presence of diabetes amongs patients.

```python
### Importing modules and loading in data

import sklearn
import numpy 
from sklearn import datasets

data = datasets.load_diabetes()
```

<!-- grade_id: cell-69d4f2db857ab81e, points: 0, grade: false, solution: true -->
```python
## Part A (20 pts)

'''
Use a train/test split to divide the data into X_train, X_test, y_train, and y_test 
correpsonding to the input features and the output vector. 
For testing purposes, use a random state of 10.
'''

from sklearn.model_selection import train_test_split

# your code here
from sklearn.model_selection import train_test_split

X = data['data']
y = data['target']

X_train, X_test, y_train, y_test = train_test_split(X, y, random_state=10)
```

<!-- grade_id: cell-95d2f03d3c521bf1, points: 20, grade: true, solution: false -->
```python

''' 
End of Section
'''
```

<!-- grade_id: cell-8233dbb26a115afd, points: 0, grade: false, solution: true -->
```python
## Part B (20 pts)

'''
Run a simple linear regresison algorithm on the input data. 
Store the fitted model in the variable 'reg'.

Hint: You may refer to the documentation on the following page: 
https://scikit-learn.org/stable/modules/generated/sklearn.linear_model.LinearRegression.html
'''

# your code here
from sklearn.linear_model import LinearRegression
reg = LinearRegression().fit(X_train, y_train)
```

<!-- grade_id: cell-43064c56ebe8afd8, points: 20, grade: true, solution: false -->
```python

''' 
End of Section
'''
```

<!-- grade_id: cell-ff0760c3016cf090, points: 0, grade: false, solution: true -->
```python
## Part C (20 pts)

'''
Run a random forest regressor on the input data. 
Store the fitted model in the variable 'rf'.
Use a max_depth of 5 and a random state of 42.

Hint: You may refer to the documentation on the following page: 
https://scikit-learn.org/stable/modules/generated/sklearn.ensemble.RandomForestRegressor.html
'''

# your code here
from sklearn.ensemble import RandomForestRegressor
rf = RandomForestRegressor(max_depth=5, random_state=42).fit(X_train, y_train)
```

<!-- grade_id: cell-3c1b9ad5a56c348e, points: 20, grade: true, solution: false -->
```python
### BEGIN HIDDEN_TESTS
assert rf.predict(X)[10] == 101.50989829460126
### END HIDDEN_TESTS

''' 
End of Section
'''
```

Output:
```
' \nEnd of Section\n'
```

<!-- grade_id: cell-44efaf77274cfbb6, points: 0, grade: false, solution: true -->
```python
## Part D (20 pts)

'''
Run the MLP regresison algorithm on the input data. 
Use a hidden layer size of 10, stochastic gradient descent solver, and a random state of 42.
Store the fitted model in the variable 'nn'. 

Hint: You may refer to the docomentation on the following page: 
https://scikit-learn.org/stable/modules/generated/sklearn.neural_network.MLPRegressor.html#sklearn.neural_network.MLPRegressor
'''

# your code here
from sklearn.neural_network import MLPRegressor
nn = MLPRegressor(hidden_layer_sizes = 10, solver = 'sgd', random_state=42).fit(X_train, y_train)
```

<!-- grade_id: cell-dfc2bb2f983e253a, points: 20, grade: true, solution: false -->
```python

''' 
End of Section
'''
```

<!-- grade_id: cell-58923a8c298ff2af, points: 0, grade: false, solution: true -->
```python
## Part E (20 pts)

'''
Let's compare the results of the three algorithms on the test set.
There are numerous ways we can compare accuracy, but for the purposes of this
exercise we're going to use the built in 'score' method from the scikit-learn library.

Create a new variable called 'best'. 
Set it equal to 0 if the linear regressor performes best,
1 if the random forest performs best, or
2 if the neural network performs best.

Hint: You may refer the 'score' method at the following page:
https://scikit-learn.org/stable/modules/generated/sklearn.linear_model.LinearRegression.html#sklearn.linear_model.LinearRegression.score
'''

# your code here
acc_nn = nn.score(X_test, y_test)
acc_reg = reg.score(X_test, y_test)
acc_rf = rf.score(X_test, y_test)
best = 0

acc_nn
```

Output:
```
0.51193392442369
```

<!-- grade_id: cell-0e70e697dc2b61eb, points: 20, grade: true, solution: false -->
```python

''' 
End of Section
'''
```